rule_files:
  - test_rules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    name: "segfault test"
    input_series:
      - series: 'blobby:cluster_disk_bytes_usage{cluster="sha_dev_1",user="a_hail"}'
        values: "1+0x1"

    alert_rule_test:
      - eval_time: 1m
        groupname: blobby.user_disk_bytes_rules_3
        alertname: BlobbyUserNearBytesQuota
        exp_alerts: []
